## SI 330 Lab 2: Numpy and Pandas

```python
### RUN THIS CELL FIRST
from multiplechoice import si330_labs_mcq
import pandas as pd
import numpy as np
```

### Part 1: Numpy and Pandas Series (25 min.)

First, refresh on what we just covered in lecture by answering the following questions. Please discuss with your group and consult notes if necessary. Try to think through them instead of simply copying and executing the code. 

```python
si330_labs_mcq(0) #Numpy Iterators, Shapes, and Dimensions
```

Output:
```
HTML(value='All of the following code snippets will yield a list with the same numbers except for one implemen…
TwoByTwoLayout(children=(HTML(value='\n        <mark><b>A.</b></mark>\n        <pre>\n            arr = np.arr…
HBox(children=(Button(description='A', style=ButtonStyle()), Button(description='B', style=ButtonStyle()), But…
HTML(value='')
```

```python
si330_labs_mcq(1) #Masking
```

Output:
```
HTML(value='Given <pre>\n                       arr = np.array([3, 3, 3, 3, 5, 5, 5])\n                       …
TwoByTwoLayout(children=(HTML(value='\n        <mark><b>A.</b></mark>\n        <pre>\n            [5, 5] \n   …
HBox(children=(Button(description='A', style=ButtonStyle()), Button(description='B', style=ButtonStyle()), But…
HTML(value='')
```

Now, create a series with 7 elements as follows:
1. last name of your choosing (string)
2. past tense verb (string)
3. number (int)
4. number with two decimals (float)
5. plural noun (string)
6. adverb (generally words ending in -ly) (string)
7. verb (string)

and share this line of code with each member of your group. Return a list containing everyone's series. 

You might know where this is going already...

<!-- grade_id: cell-fc3d4d82eec0cd46, points: 0, grade: false, solution: true -->
```python
def compiled_words():
### BEGIN SOLUTION
    words1 = pd.Series(['Brooks', 'taught',  15, 9384.29,'microwaves', 'quickly', 'sue'])
    words2 = pd.Series(['Teplovs', 'baked',  0, 129384.22,'masks', 'rapidly', 'bake'])
    words3 = pd.Series(['Erickson', 'cringed',  5, 4.12,'cheetos', 'poorly', 'rake'])
    words4 = pd.Series(['Brooks', 'yelled',  19, 89000.01,'umbrellas', 'quickly', 'cry'])
    return [words1, words2, words3, words4]
### END SOLUTION
```

<!-- grade_id: cell-d1580d7e01f67395, points: 1, grade: true, solution: false -->
```python
assert isinstance(compiled_words()[0], pd.Series)
assert len(compiled_words()) >= 4
```

Next, write a function to do the following:

Create an $n \times 7$ numpy array with one `True` in each column at a random index and `False` everywhere else (where $n$ is the number of people in your group, so 4 or 5):

   $$
   \begin{pmatrix}
    True &  F & \dots & True & F\\
    F & F & \dots & F & F\\
    F & F & \dots & F & F \\
    F & True & \dots & F & True \\
    \end{pmatrix}
   $$
   
Hint: use [np.full](https://numpy.org/doc/stable/reference/generated/numpy.full.html) (view the documentation!) to create a matrix where every element is `False`. Then, go through column by column and assign a random index to be `True`. 


<!-- grade_id: cell-b19f19d2ccfc107e, points: 0, grade: false, solution: true -->
```python
def boolean_mask_matrix():
    ### BEGIN SOLUTION
    matrix = np.full((4,7), False)
    for j in np.arange(matrix.shape[1]):
        i = np.random.randint(0, 3)
        matrix[i, j] = True
    
    return matrix
    ### END SOLUTION
```

<!-- grade_id: cell-498d3dcb429acde9, points: 1, grade: true, solution: false -->
```python
test = boolean_mask_matrix()
assert test.shape == (4, 7) or test.shape == (5, 7)
assert isinstance(test[2,3], np.bool_)
```

Now we'll use the matrix rows as a mask to create a new series, which will contains words from multiple people in your group. To do so, we will once again create a list of series (except each one will have a subset of the 7 possible words this time).

More specifically, for the first series in `word_series`, only grab the words that correspond to `True` in the first row of the mask matrix. 

![example of masks and series](part2_example.jpg)

Repeat this for each series-row pair and append the resulting *series* to `new_list`. The `pd_concat` line then combines all of these series together for you.

<!-- grade_id: cell-54ebbe1d647931d3, points: 1, grade: true, solution: true -->
```python
mask = boolean_mask_matrix()
word_series = compiled_words()

def get_final_words():
    new_list = []
    ### BEGIN SOLUTION
    for i in range(mask.shape[0]):
        new_list.append(word_series[i][mask[i,:]])
    
    ### END SOLUTION
    new_list = pd.concat(new_list)
    return new_list
```

Assuming your code above is correct, enjoy this computerized version of Mad Libs inspired by recent events! You can rerun the previous cells to get a different randomized result. 

```python
final_words = get_final_words()  
madlibs = """This semester, I have been very unimpressed with administrator %s's pandemic response.
On several different occations, I have %s at board meetings. I have also called the school %i times
to complain and tell them that my %.2f COVID fee should be spent on appeasement %s, 
but the administration has not listened. If this is not resolved %s, I will %s. 
""" % tuple(final_words[i] for i in range(len(final_words)))
print(madlibs)
```

Output:
```
This semester, I have been very unimpressed with administrator Teplovs's pandemic response.
On several different occations, I have baked at board meetings. I have also called the school 15 times
to complain and tell them that my 9384.29 COVID fee should be spent on appeasement masks, 
but the administration has not listened. If this is not resolved poorly, I will bake.
```

### Part 2: Pandas Dataframes (15 min.)

In this section, we'll be doing a classic icebreaker activity (kind of), but first, let's do a quick refresher regarding dataframes. 

```python
si330_labs_mcq(2) #Dataframe Indexing
```

Output:
```
HTML(value="Consider the following dataframe:<pre>\n                df = pd.DataFrame([{'fish': 'black bass', …
TwoByTwoLayout(children=(HTML(value="\n        <mark><b>A.</b></mark>\n        <pre>\n            df.loc['rive…
HBox(children=(Button(description='A', style=ButtonStyle()), Button(description='B', style=ButtonStyle()), But…
HTML(value='')
```

Create a dataframe with the following columns:

- Name (just first names)
- Year (referring to year of study; record as an int 1 = Freshman, ..., 4 = Senior)
- Claim 1 (either a truth or a lie)
- Claim 2
- Claim 3

Take a few minutes to populate this dataframe with the names of everyone in your group, two truths about each of you, and one lie. 

<!-- grade_id: cell-1efe84dd8c9b4887, points: 0, grade: false, solution: true -->
```python
def get_two_truths_and_a_lie():
    ### BEGIN SOLUTION
    info = [{'Name': 'John', 'Year': 3, 'Claim 1': 'I like pie.', 'Claim 2': 'I like bananas.', 'Claim 3': 'I like banana cream pie.'},
            {'Name': 'Jane', 'Year': 4, 'Claim 1': 'I like cake.', 'Claim 2': 'I like chocolate.', 'Claim 3': 'I like chocolate cake.'},
            {'Name': 'Jack', 'Year': 3, 'Claim 1': 'I like pie.', 'Claim 2': 'I like chicken.', 'Claim 3': 'I like chicken pot pie.'},
            {'Name': 'Jill', 'Year': 4, 'Claim 1': 'I like cake.', 'Claim 2': 'I like coconut.', 'Claim 3': 'I like coconut cake.'}]
    df = pd.DataFrame(info)
    ### END SOLUTION
    return df
```

<!-- grade_id: cell-ad7d85904cd9136e, points: 1, grade: true, solution: false -->
```python
assert isinstance(get_two_truths_and_a_lie(), pd.DataFrame)
```

Now let's try manipulating and indexing this dataframe.

Write a function that returns a tuple `(a, b)`:

where `a` is a dataframe from `df` called `sample_demographic_df` which only includes columns `Name` and `Year` and the first two rows.

and `b` is a dataframe from `df` called `sample_icebreaker_df` which only includes the three claims from the remaining rows.

<!-- grade_id: cell-1775ef593876b996, points: 0, grade: false, solution: true -->
```python
df = get_two_truths_and_a_lie()
def split_frame():
    ### BEGIN SOLUTION
    sample_demographic_df = df[['Name', 'Year']].iloc[0:2]
    sample_icebreaker_df = df[['Claim 1', 'Claim 2', 'Claim 3']].iloc[2:]
    ### END SOLUTION
    return (sample_demographic_df, sample_icebreaker_df)
```

<!-- grade_id: cell-38af73d8df11d9a9, points: 0, grade: true, solution: false -->
```python
x, y = split_frame()
assert x.shape == (2, 2)
assert x['Name'].shape == (2,)
assert x['Year'].shape == (2,)
assert y['Claim 2'].shape == (2,) or y['Claim 2'].shape == (3,)
```

Finally, talk to your peers to actually find out which claim is the lie! Once you do, modify the original `df` to include a column named `Lie` that contains an int, either `1`, `2`, or `3`, to indicate which claim is the lie.

<!-- grade_id: cell-0d69bd9315a44790, points: 1, grade: true, solution: true -->
```python
### BEGIN SOLUTION
df['Lie'] = [1, 2, 3, 2]
### END SOLUTION
```